## Q3
จงเขียนโปรแกรมที่อ่านไฟล์ studentScore.txt ที่มีรายชื่อนักเรียนและคะแนนสอบดังแสดงในตัวอย่างข้างล่างเก็บเป็น dictionary แล้วแสดงชื่อนักเรียนเรียงจากคะแนนสูงสุดไปต่ำสุด กำหนดให้ ไม่มีใครได้คะแนนเท่ากัน

### ตัวอย่างไฟล์ studentScore1.txt: 

Piti 25 <br>
Mana 32 <br>
Ari 19 <br>
Suda 28 <br>
Somsak 26 <br>

### ตัวอย่างไฟล์ studentScore2.txt: 

Manee 22 <br>
Piti 32 <br>
Ton 35 <br>
Tree 21 <br>
Siri 23 <br>
Mon 20 <br>

### ตัวอย่างการทำงาน: 
Choose your studentScore: **studentScore1.txt** <br>
Mana <br>
Suda <br>
Somsak <br>
Piti <br>
Ari <br>

### ตัวอย่างการทำงาน: 
Choose your studentScore: **studentScore2.txt** <br>
Ton <br>
Piti <br>
Siri <br>
Manee <br>
Tree <br>
Mon

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def get_dict(file_name):
    ''' สร้าง dictionary จากไฟล์ file_name'''
    f = open(file_name)
    sd_score = {}

    for line in f:
        name, score = line.split()
        sd_score[name] = int(score)
    
    return(sd_score)
    
def run():
    ''' รับชื่อไฟล์ที่มีรายชื่อนักเรียนและคะแนนสอบ แล้วแสดงชื่อนักเรียนเรียงจากคะแนนสูงสุดไปต่ำสุด กำหนดให้ ไม่มีใครได้คะแนนเท่ากัน'''
    file = input('Choose your studentScore:')
    sd_score = get_dict(file)
    
    score = sorted(list(sd_score.values()))
    name = []
    for i in score:
        for j in sd_score:
            if sd_score[j] == i:
                name.append(j)
                
    for i in range(-1,-(len(name)+1),-1):
        print(name[i])
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Choose your studentScore: studentScore1.txt
Mana
Suda
Somsak
Piti
Ari
```

```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO

def test(inputs,expected_output,f): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        f()
        assert ''.join(mock_stdout.getvalue().lower().split()) == ''.join(expected_output.lower().split()),mock_stdout.getvalue()
def test_easy_print(inputs,expected_output):
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert expected_output in mock_stdout.getvalue(), mock_stdout.getvalue()
```

<!-- grade_id: cell-baa6351bd29ebb54, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบการแสดงผลลัพธ์อย่างง่าย'''
test_easy_print(inputs=['studentScore1.txt'],expected_output='M')
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-159cd7c85a0c7a54, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
assert get_dict('studentScore1.txt') == {'Piti': 25, 'Mana': 32, 'Ari': 19, 'Suda': 28, 'Somsak': 26}
assert get_dict('studentScore2.txt') == {'Manee': 22, 'Piti': 32, 'Ton': 35, 'Tree': 21, 'Siri': 23, 'Mon': 20}
print('ok')
```

Output:
```
ok
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['studentScore1.txt'],'Mana\nSuda\nSomsak\nPiti\nAri\n',f=run) 
test(['studentScore2.txt'],'Ton\nPiti\nSiri\nManee\nTree\nMon\n',f=run) 
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-9241545eb2b30f20, points: 1, grade: true, solution: false -->
```python
'เทสลับ'
```

Output:
```
'เทสลับ'
```

```python

```